Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Björn Steinorth"
COLLABORATORS = "Jan Wirth"
```

---

# HASHING

Schreiben Sie eine einfache Hash-Funktion `my_hash(string, s)` mit der Tabellengröße $s$ mittels Modulo Operation (auch bekannt als *Lineares Sondieren* oder  *division remainder hashing*. (4 Punkte)

<!-- grade_id: cell-af1b70e92b19d8c4, points: 4, grade: true, solution: true -->
```python
def my_hash(string, s, i=0):#i = 0, wenn nur wie in Aufgabe 1 der Hashwert berechnet werden soll
    summe = 0;
    for j in range (0,len(string)):
        #nimmt den ASCII-Wert multipliziert mit der Stelle des characters beginnend bei 1
        #um viele verschiedene Werte zu erhalten und bildet mit allen Werten die Summe
        summe = summe + (j+1)*ord(string[j])
    #i für Verschiebung mit linearem Sondieren
    hashwert = (summe+i) % s #es darf nur s Elemente danach geben
    return hashwert
```

Berechnen Sie mittels der von Ihnen geschriebene Funhtion die Hashwerte der Zeichenketten **done** und **node** mit einer Tabellengröße von $s = 3$. (3 Punkte)

<!-- grade_id: cell-b27ef90cc2132399, points: 3, grade: true, solution: true -->
```python
hashwert_done = my_hash("done",3)
hashwert_node = my_hash("node",3)
print("Hashwert von done = ",hashwert_done)
print("Hashwert von node = ",hashwert_node)
```

Output:
```
Hashwert von done =  0
Hashwert von node =  1
```

Ist der Hash beider Zeichenketten der Gleiche? Nennen Sie den Grund. (3. Punkte)

<!-- grade_id: cell-71dfab42b9d4fe7a, points: 3, grade: true, solution: true -->
Der Hashwert ist bei dieser Implementation nicht gleiche, da die ASCII-Werte der Character mit der Stelle multipliziert werden, damit die Funktion nicht direkt durchschaubar ist. Sonst hätten alle Strings die genau die gleichen Character haben den gleichen Hashwert.

Eine weit verbreitete Methode zur Kollisionsvermeidung ist die *Lineare Suche (linear probing)*. Erzeugen Sie eine Hash-Tabelle der Größe $s = 11$ in der Form eines Assoziativen Datenfelds (Dictionary) für die Liste `lisf = [10,45,43,76,57,12,77,13]` Die Hash-Tabelle sollte wie folgt aussehen: `table = {0: 76, 1: 12, 2: 45, 3: 10, 4: 43, 5: 77, 6: 13, 7: None, 8: None, 9: None, 10: 57}`. (10 Punkte)

<!-- grade_id: cell-f34574fb999a906f, points: 0, grade: false, solution: true -->
```python
def hashvalue(item, hash_table_size, i):
    return (item % hash_table_size + i) % hash_table_size

def linearProbingHash(item_list, hash_table_size):
    #Erzeugen der Hashtabelle als Dictionary mit hash_table_size Elementen
    hash_table = {}
    for i in range (0, hash_table_size):
        hash_table[i] = None  
    #für jedes Element der item_list wird der Hashwert berechnet
    for j in range (0, len(item_list)):
        #Parameter i für Verschiebung bei Kollision
        for i in range (0,hash_table_size):
            hashwert = hashvalue(item_list[j], hash_table_size, i)
            #Hashwert wird so lange mit i neu berechnet, bis ein Key ohne Value gefunden wird
            #Das Element der Liste wird an der Stelle des berechneten Hashwertes eingetragen
            if hash_table[hashwert] == None:
                hash_table[hashwert] = item_list[j]
                break #i muss nicht weiter erhöht werden, wenn das Element zugeordnet wurde
    return hash_table

list = [10,45,43,76,57,12,77,13]
s = 11
hashtabelle = linearProbingHash(list, s)
print("Hashtabelle:",hashtabelle)
            
        
```

Output:
```
Hashtabelle: {0: 43, 1: 45, 2: 76, 3: 57, 4: 12, 5: 77, 6: 13, 7: None, 8: None, 9: None, 10: 10}
```

<!-- grade_id: cell-9d77dc861e6c0eb0, points: 10, grade: true, solution: false -->
```python
list = [10,45,43,76,57,12,77,13]

```

Erzeugen Sie eine Hash-Tabelle der Größe $s = 11$ in Form eines Assoziativen Datenfelds für die Liste `list = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"].`Die Hash-Tabelle sollte aussehen wie folgt: `table = {0: 'her', 1: 'this', ... , 10: None}`. (10 Punkte)

<!-- grade_id: cell-edc469c308109c3d, points: 0, grade: false, solution: true -->
```python
def linearProbingHash(item_list, hash_table_size):
    #Erzeugen der Hashtabelle als Dictionary mit hash_table_size Elementen
    hash_table = {}
    for i in range (0, hash_table_size):
        hash_table[i] = None  
    #für jedes Element der item_list wird der Hashwert berechnet
    for j in range (0, len(item_list)):
        #Parameter i für Verschiebung bei Kollision
        for i in range (0,hash_table_size):
            #verwendung der my_hash Funktion von oben
            hashwert = my_hash(item_list[j], hash_table_size, i)
            if hash_table[hashwert] == None:
                hash_table[hashwert] = item_list[j]
                break        
    return hash_table

item_list = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"]
hash_table_size = 11
hashtabelle = linearProbingHashStrings(item_list, hash_table_size)
print("Hashtabelle:",hashtabelle)    

```

Output:
```
Hashtabelle: {0: 'her', 1: 'this', 2: 'that', 3: 'when', 4: 'what', 5: 'how', 6: 'i dont know', 7: None, 8: 'why', 9: None, 10: 'his'}
```

<!-- grade_id: cell-df52a49cc705fd10, points: 10, grade: true, solution: false -->
```python
stringList = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"]
```

Schreiben Sie eine Funktien `all_cocktails(filename)`, die die Datei `cocktails.json` in ein Assozietives Datenfeld `recipes` liest und schreiben Sie eine Funktion `all_ingredients(recipes)`, die eine komplette Liste aller Zutaten ausgibt. (10 Punkte)

<!-- grade_id: cell-a39e1b1cae7075df, points: 0, grade: false, solution: true -->
```python
import json

def all_cocktails(filename):
    with open(filename) as f:
        return json.load(f)
        
def all_ingredients(recipes):
    ingredients = set()
    for c in recipes['cocktails']:
        for i in c['ingredients']:
            try:
                ingredients.add(i['ingredient'])
            except KeyError:
                pass
        
    return ingredients

recipes = all_cocktails('cocktails.json')
numberOfIngredients = len(all_ingredients(recipes))
print(numberOfIngredients)
print (all_ingredients(recipes))
print (numberOfIngredients)

```

<!-- grade_id: cell-4d70aaa8919259a4, points: 10, grade: true, solution: false -->
```python
recipes = all_cocktails('cocktails.json')

assert (numberOfIngredients) == (37)
assert ('Apricot brandy' in all_ingredients(recipes)) 
assert ('Pineapple juice' in all_ingredients(recipes))
assert ('Campari' in all_ingredients(recipes)) 
assert ('Kirsch' in all_ingredients(recipes))
assert ('Pisco' in all_ingredients(recipes))
```