Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# HASHING

Schreiben Sie eine einfache Hash-Funktion `my_hash(string, s)` mit der Tabellengröße $s$ mittels Modulo Operation (auch bekannt als *Lineares Sondieren* oder  *division remainder hashing*. (4 Punkte)

<!-- grade_id: cell-af1b70e92b19d8c4, points: 4, grade: true, solution: true -->
```python
def my_hash(string, s):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Berechnen Sie mittels der von Ihnen geschriebene Funhtion die Hashwerte der Zeichenketten **done** und **node** mit einer Tabellengröße von $s = 3$. (3 Punkte)

<!-- grade_id: cell-b27ef90cc2132399, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Ist der Hash beider Zeichenketten der Gleiche? Nennen Sie den Grund. (3. Punkte)

<!-- grade_id: cell-71dfab42b9d4fe7a, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE

Eine weit verbreitete Methode zur Kollisionsvermeidung ist die *Lineare Suche (linear probing)*. Erzeugen Sie eine Hash-Tabelle der Größe $s = 11$ in der Form eines Assoziativen Datenfelds (Dictionary) für die Liste `lisf = [10,45,43,76,57,12,77,13]` Die Hash-Tabelle sollte wie folgt aussehen: `table = {0: 76, 1: 12, 2: 45, 3: 10, 4: 43, 5: 77, 6: 13, 7: None, 8: None, 9: None, 10: 57}`. (10 Punkte)

<!-- grade_id: cell-f34574fb999a906f, points: 0, grade: false, solution: true -->
```python
def linearProbingHash(item_list, hash_table_size):
    
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9d77dc861e6c0eb0, points: 10, grade: true, solution: false -->
```python
list = [10,45,43,76,57,12,77,13]

```

Erzeugen Sie eine Hash-Tabelle der Größe $s = 11$ in Form eines Assoziativen Datenfelds für die Liste `list = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"].`Die Hash-Tabelle sollte aussehen wie folgt: `table = {0: 'her', 1: 'this', ... , 10: None}`. (10 Punkte)

<!-- grade_id: cell-edc469c308109c3d, points: 0, grade: false, solution: true -->
```python
def linearProbingHashStrings(item_list, hash_table_size):
    
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-df52a49cc705fd10, points: 10, grade: true, solution: false -->
```python
stringList = ["his", "her", "this", "that", "what", "when", "how", "why", "i dont know"]
```

Schreiben Sie eine Funktien `all_cocktails(filename)`, die die Datei `cocktails.json` in ein Assozietives Datenfeld `recipes` liest und schreiben Sie eine Funktion `all_ingredients(recipes)`, die eine komplette Liste aller Zutaten ausgibt. (10 Punkte)

<!-- grade_id: cell-a39e1b1cae7075df, points: 0, grade: false, solution: true -->
```python
import json

def all_cocktails(filename):
    with open(filename) as f:
        return json.load(f)
        
def all_ingredients(recipes):
    ingredients = set()
    for c in recipes['cocktails']:
        for i in c['ingredients']:
            try:
                ingredients.add(i['ingredient'])
            except KeyError:
                pass
        
    return ingredients

recipes = all_cocktails('cocktails.json')
numberOfIngredients = len(all_ingredients(recipes))
print(numberOfIngredients)
print (all_ingredients(recipes))
print (numberOfIngredients)

```

Output:
```
37
{'Calvados', 'Cream', 'Lime juice', 'Champagne', 'Campari', 'Creme liqueur', 'Drambuie', 'Soda water', 'Whiskey', 'Dark rum', 'Cranberry juice', 'Cherry liqueur', 'Syrup', 'Red Port', 'Cognac', 'Vodka', 'Triple Sec', 'Cola', 'Pisco', 'Coffee liqueur', 'Prosecco', 'Gin', 'Kirsch', 'Egg yolk', 'White rum', 'Tequila', 'Apricot brandy', 'Grapefruit juice', 'Pineapple juice', 'Coconut milk', 'Orange juice', 'Vermouth', 'Lemon juice', 'Peach puree', 'Ginger Ale', 'Galliano', 'Lillet Blonde'}
37
```

<!-- grade_id: cell-4d70aaa8919259a4, points: 10, grade: true, solution: false -->
```python
recipes = all_cocktails('cocktails.json')

assert (numberOfIngredients) == (37)
assert ('Apricot brandy' in all_ingredients(recipes)) 
assert ('Pineapple juice' in all_ingredients(recipes))
assert ('Campari' in all_ingredients(recipes)) 
assert ('Kirsch' in all_ingredients(recipes))
assert ('Pisco' in all_ingredients(recipes))
```

```python

```

```python

```

```python

```