Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jan Wirth"
COLLABORATORS = "Björn Steinorth"
```

---

# Aufgabe 1 - Binäre Suchbäume (16 Punkte)

a)	**(5 Punkte)**  
In der Vorlesung haben wir die Klasse SearchTree behandelt, die (key, value)-Paare in Node-Objekten mit folgender Definition ablegt:

```python
class Node:
    def __init__(self, key, value):
        self.key = key
        self.value = value
        self.left = self.right = None
```

Das Grundgerüst der Klasse SearchTree hat folgende Form:

<!-- grade_id: task1a-class2, points: 0, grade: false, solution: true -->
```python
def build(nodes):
    t = SearchTree();
    for node in nodes:
        t.insert(node)
    return t.root

def insert(node, key, value):
    if node.key < key:
        if node.right == None or node.right.key == key:
            node.right = Node(key, value)
        else:
            insert(node.right, key, value)
    else:
        if node.left == None or node.left.key == key:
            node.left = Node(key, value)
        else:
            insert(node.left, key, value)

def rebuildFromChildren(node):
    return build(getChildren(node))

def remove(node, key):
    if node.key < key:
        if node.right == None:
            raise KeyError('Key ' + str(key) + ' not found in tree')
        if node.right.key == key:
            node.right = rebuildFromChildren(node.right)
        else:
            remove(node.right, key)
    else:
        if node.left == None:
            raise KeyError('Key ' + str(key) + ' not found in tree')
        if node.left.key == key:
            node.left = rebuildFromChildren(node.left)
        else:
            remove(node.left, key)

def getChildren(node):
    children = []
    if node.left != None:
        children = children + getChildren(node.left)
        children.append(node.left)
    if node.right != None:
        children = children + getChildren(node.right)
        children.append(node.right)
    return children

def find(node, key):
    if node == None or node.key == key:
        return node
    elif node.key < key:
        return find(node.right, key)
    else:
        return find(node.left, key)

class SearchTree:
    def __init__(self):
        self.root = None
        self.size = 0

    def __len__(self):
        return self.size

    def insert(self, key, value):
        # YOUR CODE HERE
        if(self.root == None):
            self.root = Node(key, value)
        else:
            insert(self.root, key, value)

        
    def remove(self, key):
        if self.root == None:
            raise KeyError('Key ' + str(key) + ' not found in tree')
        elif self.root.key == key:
            self.root = rebuildFromChildren(self.root)
        else:
            remove(self.root, key)
        
    def find(self, key):
        return find(self.root, key)
        
# Traverse
def preOrderTraverse(root):
    if root is not None:
        print ('key', root.key, 'value', root.value)
        preOrderTraverse(root.left)
        preOrderTraverse(root.right)
            
def inOrderTraverse(root):
    if root is not None:
        inOrderTraverse(root.left)
        print ('key', root.key, 'value', root.value)
        inOrderTraverse(root.right)

def postOrderTraverse(root):
    if root is not None:
        postOrderTraverse(root.left)
        postOrderTraverse(root.right)
        print ('key', root.key, 'value', root.value)

```

**Hint:**  
Geben Sie einer vollständigen Implementation ab. Wenn tree ein Objekt vom Typ SearchTree ist, soll die Semantik der Aufrufe wie folgt realisiert werden:  

>*tree.insert(key, value):*   
Fügt den gegebenen Wert value unter dem Schlüssel key ein. Falls key schon vorhanden war, soll der zuvor gespeicherte Wert überschrieben werden. Andernfalls soll ein neuer Node so eingefügt werden, dass die Suchbaumbedingung erhalten bleibt.  

>*tree.remove(key):*   
Löscht den Node, der den gegebenen Schlüssel enthält, wobei die Such-baumbedingung erhalten bleibt. Falls der Schlüssel nicht vorhanden ist, soll eine KeyError-Excpetion ausgelöst werden, deren Fehlermeldung den Schlüssel angibt.  

>*found=tree.find(key):*   
Gibt den Node zurück, der den Schlüssel key enthält, oder None, wenn der Schlüssel nicht im Baum vorhanden ist.


<!-- grade_id: task1a-test, points: 5, grade: true, solution: false -->
```python
a = SearchTree();
assert a.__len__() == 0
a.insert(6,60) 
a.insert(2,20)
a.insert(5,50)
a.insert(3,30)
a.insert(8,80)
a.insert(4,40)
a.insert(1,10)
a.insert(7,70)
assert a.root.left.right.left.right.key == 4
assert a.root.left.right.left.right.value == 40
assert a.root.right.left.key == 7
assert a.root.right.left.value == 70
node1 = a.find(5)
assert node1.key == 5
assert node1.value == 50
a.insert(5,500)
node1 = a.find(5)
assert node1.key == 5
assert node1.value == 500
a.remove(5)
assert a.find(5) == None

print('preOrderTraverse:')
preOrderTraverse(a.root)
print('inOrderTraverse:')
inOrderTraverse(a.root)
print('postOrderTraverse:')
postOrderTraverse(a.root)
    
```

Output:
```
preOrderTraverse:
key 6 value 60
key 2 value 20
key 1 value 10
key 8 value 80
key 7 value 70
inOrderTraverse:
key 1 value 10
key 2 value 20
key 6 value 60
key 7 value 70
key 8 value 80
postOrderTraverse:
key 1 value 10
key 2 value 20
key 7 value 70
key 8 value 80
key 6 value 60
```

b)	**(4 Punkte)**  
Entwickeln Sie einen Algorithmus, der die Tiefe des Baumes (den maximalen Abstand von der Wurzel zu einem Blatt) bestimmt und implementieren Sie ihn als Methode, so dass depth=tree.depth() die Tiefe zurückgibt.   
**Hint:** If there is only one node, its depth is 0. 

<!-- grade_id: task1b-answer, points: 0, grade: false, solution: true -->
```python
def nodeDepth(node, currentDepth = -1):
    if node == None:
        return currentDepth
    else:
        return max(nodeDepth(node.left, currentDepth + 1),
                   nodeDepth(node.right, currentDepth + 1))
def depth(tree):
    return nodeDepth(tree.root)
    
SearchTree.depth = depth

```

<!-- grade_id: task1b-test, points: 4, grade: true, solution: false -->
```python
b = SearchTree();
assert b.depth() == -1
b.insert(6,60)
assert b.depth() == 0
b.insert(2,20)
assert b.depth() == 1
b.insert(5,50)
assert b.depth() == 2
b.insert(3,30)
assert b.depth() == 3
b.insert(8,80)
assert b.depth() == 3
b.insert(4,40)
assert b.depth() == 4
b.insert(1,10)
assert b.depth() == 4
b.insert(7,70)
assert b.depth() == 4

```

c)	**(4 Punkte)**   
Angenommen, Sie können die Schlüssel in einer selbstgewählten Reihenfolge einfügen. Welche Reihenfolge wählen Sie, damit der Baum nach dem Einfügen eine möglichst geringe Tiefe hat?

<!-- grade_id: task1c-answer, points: 4, grade: true, solution: true -->
Wie bei Divide and conquer:
1. liste sortieren
2. median als pivot wählen und in den tree einfügen
3. median pivot der linken liste einfügen
4. " rechten "
5. ganz-linke (erste von vier) listen pivotieren und wiederholen, dann die zweite etc.

d)	**(3 Punkte)**  
Beweisen oder widerlegen Sie folgende Aussage: Wenn aus einem Suchbaum erst der Schlüssel X und danach der Schlüssel Y entfernt wird, entsteht der gleiche Baum wie bei umgekehrter Reihenfolge.

<!-- grade_id: task1d-answer, points: 3, grade: true, solution: true -->
Der selbe Baum wird aufgebaut.

Da der subtree (in unserer implementierung) von links nach rechts dekonstruiert und wieder aufgebaut wird, ist der vorgang deterministisch und führt zum gleichen resultat. Siehe Skizze.
remove(5); remove(3)
```
    5             3          2
 3     7  (-5)> 2   4  (-3)>   4
2 4   6 8            7          7 
                    6 8        6 8
```
remove(3); remove(5)
```
    5              5            2
 3     7  (-3)> 2     7  (-5)>   4
2 4   6 8        4   6 8          7 
                                 6 8
```

# Aufgabe 2 – Taschenrechner (24 Punkte)

Binärbäume eignen sich auch, um mathematische Ausdrücke auszuwerten, die als Zeichenketten der Form "2+5\*3" oder "2\*4\*(3+(4-7)\*8)-(1-6)" gegeben sind. Man bezeichnet solche Bäume als *Syntax-bäume* (https://de.wikipedia.org/wiki/Abstrakter_Syntaxbaum).  

Die Ausdrücke können die arithmetischen Operationen +, -, \*, / mit den üblichen Rechenregeln enthalten (Klammern haben die höchste Priorität, Punktrechnung geht vor Strichrechnung). Zahlen sollen der Einfachheit halber immer einstellig und positiv sein. Variablen und Funktionen kommen nicht vor. Operationen mit gleicher Priorität werden von links nach rechts ausgewertet (sogenannte *Links-Assoziativität*), damit wie gewohnt 5-2+3 = (5-2)+3 = 6 gilt (und nicht 5-(2+3) = 0, was bei Auswertung von rechts herauskäme). Steht jedoch rechts von einer Zahl oder einem Klammerausdruck eine Operation mit höherer Priorität als links, wird der rechte Operator zuerst ausgewertet. Trifft man auf eine öffnende Klammer, muss man den Substring bis zur zugehörigen schließenden Klammer suchen und die Auswertung rekursiv auf diesen Substring anwenden. Dadurch ergibt sich ein Binärbaum.  

a)	**(14 Punkte)**   
Entwickeln Sie einen Algorithmus, der den zu einem Ausdruck korrespondierenden Binärbaum aufbaut, wobei jeder innere Knoten einen Operator (+, -, \*, /) repräsentiert, jeder Unterbaum einen linken bzw. rechten Operanden, und jedes Blatt eine Zahl. Der Baum soll dann durch eine Funktion parse(s) erstellt werden, an die der Ausdruck als String übergeben wird und die den Wurzelknoten des Baums zurückgibt (die Verwendung der Python-Funktionen eval() bzw. exec ist dabei *nicht* erlaubt). Implementieren Sie diese Funktion und erklären Sie in Kommentaren, wie der Algorithmus vorgeht.   
**Hinweis:** Implementieren Sie zunächst zwei Klassen Number und Operator, die als Blattknoten bzw. innere Knoten des Syntaxbaums dienen. Number-Objekte speichern also eine Zahl, und Operator-Objekte ein Operatorsymbol sowie den linken und rechten Operanden (Unterbaum). 


<!-- grade_id: task2a-code, points: 0, grade: false, solution: true -->
```python
#!/usr/bin/python
# -*- coding: utf-8 -*-
import re

OPERATORS = '|\+|\-|\*|\/'
NUMBERS = '\(?[0-9]\)?'

OPERATOR_FNS = {
    '+': lambda x, y: x + y,
    '-': lambda x, y: x - y,
    '*': lambda x, y: x * y,
    '/': lambda x, y: x / y,
    }


class NodeNum:

    def __init__(self, value):
        v = int(value)
        self.value = v

    def eval(self):
        return self.value


class NodeOp:

    def __init__(self, value):
        self.value = value
        self.left = self.right = None

    def eval(self):
        return OPERATOR_FNS[self.value](self.left.eval(),
                self.right.eval())


def groupRest(restParts):
    opening = 0
    closing = 0
    for (index, part) in enumerate(restParts):
        if '(' in part:
            opening = opening + 1
        elif ')' in part:
            closing = closing + 1
        if opening == closing:
            innerWithBrackets = ' '.join(restParts[0:index + 1])
            inner = innerWithBrackets[1:len(innerWithBrackets) - 1]
            break
    res = [split(inner)] + split(' '.join(restParts[index
                                 + 1:len(restParts)]))
    return res


def group(parts):
    res = []
    for (index, part) in enumerate(parts):
        if '(' in part:
            res = res + groupRest(parts[index:len(parts)])
            break
        else:
            res.append(parts[index])
    return res


def split(s):
    """
    
    """

    parts = re.findall('(' + NUMBERS + OPERATORS + ')', s)
    res = group(parts)
    return res


def build(nested):

  # flattwen double-nested list

    if len(nested) == 1 and isinstance(nested, list) \
        and isinstance(nested[0], list):
        nested = nested[0]
    rootOpPos = findRootOperatorPosition(nested)

  # if an operator was found, recurse

    if rootOpPos:
        root = NodeOp(nested[rootOpPos])
        root.left = build(nested[0:rootOpPos])
        root.right = build(nested[rootOpPos + 1:len(nested)])
        return root
    else:
        return NodeNum(nested[0])


def findRootOperatorPosition(nested):

  # we reverse the list because the leftmost operator has the highest precedence and we can't iterate from right to left

    rev = list(reversed(nested))

  # we first want the lesser-binding leftmost operators (+, -)
  # multiplication and division bind stronger, so we check their positions first and overwrite them with plus/minus positions if found

    leftMostLeastBindingOperatorPosition = None
    for (index, part) in enumerate(rev):
        if not isinstance(part, list) and ('*' in part or '/' in part):
            leftMostLeastBindingOperatorPosition = index
    for (index, part) in enumerate(rev):
        if not isinstance(part, list) and ('+' in part or '-' in part):
            leftMostLeastBindingOperatorPosition = index
    if leftMostLeastBindingOperatorPosition:

    # reverse index back

        return len(nested) - leftMostLeastBindingOperatorPosition - 1
    else:
        return None


def parse(s):
    print ('parsing', s)
    nested = split(s)
    tree = build(nested)
    return tree


# Traverse

def preOrderTraverse2(node):
    if type(node) == NodeOp:
        print ('key', node.key, 'value', node.value)
        preOrderTraverse2(node.left)
        preOrderTraverse2(node.right)
    elif type(node) == NodeNum:
        print ('key', node.key, 'value', node.value)


def inOrderTraverse2(node):
    if type(node) == NodeOp:
        inOrderTraverse2(node.left)
        print ('key', node.key, 'value', node.value)
        inOrderTraverse2(node.right)
    elif type(node) == NodeNum:
        print ('key', node.key, 'value', node.value)


def postOrderTraverse2(node):
    if type(node) == NodeOp:
        postOrderTraverse2(node.left)
        postOrderTraverse2(node.right)
        print ('key', node.key, 'value', node.value)
    elif type(node) == NodeNum:
        print ('key', node.key, 'value', node.value)

```

<!-- grade_id: task2a-test, points: 12, grade: true, solution: true -->
```python

```

b)	**(3 Punkte)**   
Skizzieren Sie die Bäume, die sich für die Ausdrücke "2+5\*3" und "2\*4\*(3+(4-7)\*8)-(1-6)" ergeben.

<!-- grade_id: task2b-answer, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: task2b-test, points: 1, grade: true, solution: true -->
```python
root1 = NodeOp("+")
root1.left = NodeNum(2)
root1.right = NodeOp("*")
root1.right.left = NodeNum(5)
root1.right.right = NodeNum(3)
print(root1.eval())

root2 = NodeOp("-")

root2.right = NodeOp("-")
root2.right.left = NodeNum(1)
root2.right.right = NodeNum(6)

root2.left = NodeOp("*")
root2.left.left = NodeNum(2)
root2.left.right = NodeOp("*")
root2.left.right.left = NodeNum(4)
root2.left.right.right = NodeOp("*")

inner2 = NodeOp("+")
inner2.left = NodeNum(3)
inner2.right = NodeOp("*")
inner2.right.left = NodeOp("-")
inner2.right.right = NodeNum("8")
inner2.right.left.left = NodeNum(4)
inner2.right.left.right = NodeNum(7)

root2.left.right.right = inner2

print(root2.eval())


```

Output:
```
17
-163
```

c)	**(3 Punkte)**   
Implementieren Sie eine Funktion evaluateTree(root), um einen Ausdruck mit Hilfe des in a) erstellten Baums auszurechnen.

<!-- grade_id: task2c-answer, points: 0, grade: false, solution: true -->
```python
def evaluateTree(root):
    """
    This function evaluates the result of an expression given its tree root. 

    Arguments:
    root -- the root node of the tree

    Returns:
    root.value -- the result of the expression
    """
    return root.eval()
```

<!-- grade_id: task2c-test, points: 3, grade: true, solution: false -->
```python
assert evaluateTree(parse('2+5*3')) == 17
assert evaluateTree(parse('2*4*(3+(4-7)*8)-(1-6)')) == -163

```

Output:
```
parsing 2+5*3
parsing 2*4*(3+(4-7)*8)-(1-6)
```

d)	**(4 Punkte)**   
Schreiben Sie Unit Tests für Ihr Verfahren. Beachten Sie dabei die Hinweise zum Erstellen guter Tests im Kapitel "Korrektheit" des Skripts. Beispielsweise müssen Sie verschiedene Varianten der Operator-Präzedenz und Klammerung testen.

<!-- grade_id: task2d-test, points: 4, grade: true, solution: true -->
```python
assert evaluateTree(parse('2+5*3')) == 17
assert evaluateTree(parse('2+(5-3)')) == 4
assert evaluateTree(parse('2*2*2*2')) == 16
assert evaluateTree(parse('2*4*(3+(4-7)*8)-(1-6)')) == -163
assert evaluateTree(parse('(1-6)+2*4*(3+(4-7)*8)')) == -173
assert evaluateTree(parse('(1-6+2*4*3+(4-7)*8)')) == -5
```

Output:
```
parsing 2+5*3
parsing 2+(5-3)
parsing 2*2*2*2
parsing 2*4*(3+(4-7)*8)-(1-6)
parsing (1-6)+2*4*(3+(4-7)*8)
parsing (1-6+2*4*3+(4-7)*8)
```

```python

```