Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

#  Einleitung
In dieser Aufgabe soll Ihr Verständnis von Sortieralgorithmen gestärkt werden. Stellen Sie sicher, dass sie das Archiv samt Test-Dateien heruntergeladen und in einen Unterordner `test_cases` extrahiert haben (Quelle: http://blog.quibb.org/2008/11/sort-optimization/). Öffnen Sie den Ordner und schauen Sie sich den Schema an, nachdem die Dateien benannt sind.




# Sortierverfahren implementieren (6⋅4 Pkt.)
Implementieren Sie folgende Sortieralgorithmen.

- Selectionsort
- Insertionsort
- Quicksort
- Bubblesort
- Mergesort
- Heapsort

Die Sortieralgorithmen sollen eine sortierte Kopie der ursprünglichen Liste zurück geben und das Original nicht verändern (damit das Verhalten kongruent zu `sorted()`) ist.


<!-- grade_id: cell-6d812715dbe509ce, points: 0, grade: false, solution: true -->
```python
# -*- coding: utf-8 -*-

def selectionsort(data):
    for i in range (0,len(data)):
        kleinstes = data[i]
        for j in range (i, len(data)):
            if data[j] <= kleinstes:
                kleinstes = data[j]
                stelle = j
        temp = data[i]
        data[i] = data[stelle]
        data[stelle] = temp
    return data

def insertionsort(data):
    for i in range(1,len(data)):
        currentvalue = data[i]
        position = i
        while position>0 and data[position-1]>currentvalue:
            data[position]=data[position-1]
            position = position-1
        data[position]=currentvalue
    return data

def bubblesort(data):
    for i in range(0,len(data)-1):
        for j in range(0,len(data)-1-i):
            if data[j] >= data[j+1]:
                temp = data[j]
                data[j] = data[j+1]
                data[j+1] = temp
    return data
    

def quicksort(data):
    if len(data) < 1:
        return data
    pivot = data[0]
    left = []
    right = []
    for i in range (1,len(data)):
        if data[i] <= pivot:
            left.append(data[i])
        else:
            right.append(data[i])
    left = quicksort(left)
    right = quicksort(right)
    pivotelement = [pivot]
    return left + pivotelement + right 
    
def heapsort(data):
    # YOUR CODE HERE
    raise NotImplementedError()

def mergesort(data):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-848f5caf61f83589, points: 0, grade: true, solution: false -->
```python
import random

SAMPLE_LIST = random.sample(range(100), 30)

def test_algorithm(algorithm):
    sample_list = SAMPLE_LIST[:]
    assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
    assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)
```

<!-- grade_id: cell-87b11519e17f1c63, points: 4, grade: true, solution: false -->
```python
test_algorithm(selectionsort)
```

Output:
```
<ipython-input-3-aac1f048cc98> in <module>()
----> 1 test_algorithm(selectionsort)

<ipython-input-2-03b62a4e4c59> in test_algorithm(algorithm)
      6     sample_list = SAMPLE_LIST[:]
      7     assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
----> 8     assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)

AssertionError: selectionsort altered the original
```

<!-- grade_id: cell-24cfee0a571db291, points: 4, grade: true, solution: false -->
```python
test_algorithm(insertionsort)
```

Output:
```
<ipython-input-4-c42301acb0d5> in <module>()
----> 1 test_algorithm(insertionsort)

<ipython-input-2-03b62a4e4c59> in test_algorithm(algorithm)
      6     sample_list = SAMPLE_LIST[:]
      7     assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
----> 8     assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)

AssertionError: insertionsort altered the original
```

<!-- grade_id: cell-78ffe6f52e47fc47, points: 4, grade: true, solution: false -->
```python
test_algorithm(bubblesort)
```

Output:
```
<ipython-input-5-0e4804a5fe9e> in <module>()
----> 1 test_algorithm(bubblesort)

<ipython-input-2-03b62a4e4c59> in test_algorithm(algorithm)
      6     sample_list = SAMPLE_LIST[:]
      7     assert algorithm(sample_list) == sorted(sample_list), "{} did not sort the list".format(algorithm.__name__)
----> 8     assert sample_list == SAMPLE_LIST, "{} altered the original".format(algorithm.__name__)

AssertionError: bubblesort altered the original
```

<!-- grade_id: cell-e60d068bfdb9cdde, points: 4, grade: true, solution: false -->
```python
test_algorithm(quicksort)
```

<!-- grade_id: cell-037297fb44664ec1, points: 4, grade: true, solution: false -->
```python
test_algorithm(heapsort)
```

<!-- grade_id: cell-96e6149464b940df, points: 4, grade: true, solution: false -->
```python
test_algorithm(mergesort)
```

# Einlesen von Dateien (4 Pkt.)
Schreiben sie eine Funktion `load_test_cases(path)`, die zu den zur Verfügung gestellten Testfällen jeweils die Datei mit 10e3 Einträgen lädt.

<!-- grade_id: cell-d64d799097780639, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

TEST_CASE_PATH = "test_files"

def load_test_cases(path):
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-6b64b68014341d9a, points: 4, grade: true, solution: false -->
```python
import unittest

unittest_test_case = unittest.TestCase('__init__')

loaded_test_cases = load_test_cases(TEST_CASE_PATH)
available_test_cases = ['1245.repeat', 'end.0', 'identical', 'med.3.killer', 'rand.dups', 'rand.no.dups', 'rand.steps', 
                        'rev.ends', 'rev.partial', 'rev.saw', 'reverse', 'seq.0.is.1000', 'seq.partial', 'seq.saw', 'sequential']

unittest_test_case.assertCountEqual(loaded_test_cases.keys(), available_test_cases)

for key, value in loaded_test_cases.items():
    assert len(value) == 10e3, "Number of testcases mismatch for {}".format(key)
```

# Sortierverfahren testen (12 Pkt.)
 Vergleichen Sie die Sortierverfahren untereinander und mit der Python Funktion `sorted()`
 - Benutzen sie `result = %timeit -n 100 -r 3 -o` um die Ausführung zu messen und das Ergebnis in einer Variablen speichern zu können. Auf das beste Ergebnis können Sie mit `result.best` zugreifen.
 - Gruppieren Sie die Ergebnisse nach Testfall.
 - Testen sie folgende „Splices“ Ihres Testdatensatzes: `x_values = [10, 20, 50, 100, 200, 500, 1000, 2000, 5000, 10000]`
 - Stellen sie die Ergebnisse mittels Tabelle und Graphen mit Titel und Achsenbeschriftung dar (Matplotlib).
 - Bewerten Sie das Ergebnis

<!-- grade_id: cell-584a475079bfa4bd, points: 12, grade: true, solution: true -->
```python
# for quicksort
import sys
sys.setrecursionlimit(20000)

# YOUR CODE HERE
raise NotImplementedError()
        
```